<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Skriptsprachen
### Winterersemester 2022/23
Prof. Dr. Heiner Giefers

```python
import sys
try:
    import chess
except:
    !{sys.executable} -m pip install chess
```

### Das n-Damen Problem

In dieser Aufgabe betrachten wir das bekannte *n-Damen Problem*. Dabei geht es darum, auf einem Schachbrett der Größe $n\times n$ insgesamt $n$ Damen si zu plazieren, dass kein der Damen eine andere angreift. Im Schach kann sich die Dame horizontal, vertikal und diagonal bewegen. Es kann also pro Zeile und Spalte nur eine Dame palziert werden, gleichzeitig darf auf den ansteigenden und absteigenden Diagonalen nur jeweils eine Dame platziert sein.

Das Folgende Beispiel zeigt eine Lösung des 8-Damen Problems für ein normales Schachbrett:

```python
fen = "Q7/4Q3/7Q/5Q2/2Q5/6Q1/1Q6/3Q4"
board = chess.Board(fen)
display(board)
```

#### Vorgehensweise
Das n-Damen Probelm ist ein Fall, den man ideal mit dem *Backtracking-Algorithmus* lösen kann. Beim Backtracking wird versucht, eine Teillösung zu einer Komplettlösung auszubauen. Scheitert der Algorithmus auf seinem Weg zur lösung, geht er einen Schritt zurück, ändert eine Entscheidung und probiert diese neue Teillösung zur Komplettlösung zu erweitern. Auf diesem Wege (der im Prinzip einer Tiefensuche entspricht) können alle möglichen Lösungen gefunden werden.

**Lösungsideen:**
- Auf jeder Zeile des Spielbretts kann nur eine Dame plaziert werden. Wir können also in unsere Backtracking routine alle Spalten einer Zeile durchprobieren; falls eine Dame auf einer Spalte der Zeile fehlerfrei plaziert werden kann, merken wir uns die Spalte und probieren dann eine Weitere Dame auf der nächsten Zeile zu plazieren. Dies führt zu einem rekursiven Aufruf der Backtracking-Funktion.
- Immer, wenn wir auf der untersten Zeile des Feldes ankommen und dort eine Dame Plazieren können, merken wir uns die aktuelle Lösung und beenden die Backtracking-Funktion.
- Wenn wir von einem rekursiven Aufruf der Backtracking-Funktion zurückkehren, müssen wir das Plazieren der letzen Dame rückgängig machen und zum nächsten Feld der Zeile gehen. 

- Da wir über die Zeilen gehen, müssen wir uns die Zeilen, auf denen Damen plaziert worden sind nicht merken
- Die Spalte, in der wir die Damen plaziert haben, speichern wir in einer Liste. Diese Liste ist dann auch das Ergebnis für eine Lösung, falls die List *n* Spalten-IDs für die *n* plazierten Damen enthält.
- Im Gegensatz zu den Zeilen und Spalten, lassen sich die Diagonalen nicht so einfach durchnummerieren. Man kann aber beobachten, dass die unterschidlichen Diagonalen immer den gleichen Wert für `Zeile+Spalte` (ansteigende Diagonalen), bzw. `Zeile-Spalte` (absteigende Diagonalen) haben. Mit diesen Berechnungen kann man also festhalten, ob eine bestimmte ansteigende oder absteigende Diagonale bereits besetzt ist.

![](https://media.geeksforgeeks.org/wp-content/uploads/20200716201927/matrix.png)

<!-- grade_id: cell-7d81e9746387e944, points: 0, grade: false, solution: true -->
```python
class nQueens:
    
    def __init__(self, n):
        # YOUR CODE HERE
        raise NotImplementedError()
    
    def solve(self):
        n = self.n
        spalten   = []  # 
        anstDiag  = []  # (Zeile + Spalte)
        abstDiag  = []  # (Zeile - Spalte)
        
        
        def backtracking(zeile):
            # YOUR CODE HERE
            raise NotImplementedError()

        backtracking(0)
        return self.res
    
    def __str__(self):
        # YOUR CODE HERE
        raise NotImplementedError()
```

<!-- grade_id: cell-414563331e7155b6, points: 0, grade: true, solution: false -->
```python
s = nQueens(4)
assert s.solve()==[(1, 3, 0, 2), (2, 0, 3, 1)]
```

<!-- grade_id: cell-4f1a2555dc94d66b, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```